---
title: Template Chapter
description: This is a template chapter.

--- type:NormalExercise lang:r xp: skills: key:dfc165a539
## Template exercise 
This is a template exercise.aa

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
